# Association rule mining

In this notebook, you'll implement the basic pairwise association rule mining algorithm.

To keep the implementation simple, you will apply your implementation to a simplified dataset, namely, letters ("items") in words ("receipts" or "baskets").

## Problem definition

Let's say you have a fragment of text in some language. You wish to know whether there are association rules among the letters that appear in a word. In this problem:

- Words are "receipts"
- Letters within a word are "items"

You want to know whether there are _association rules_ of the form, $a \implies b$, where $a$ and $b$ are letters. You will write code to do that by calculating for each rule its _confidence_, $\mathrm{conf}(a \implies b)$, which is an estimate of the conditional probability of $b$ given $a$, or $\mathrm{Pr}[b \,|\, a]$.

## Sample text input

Let's carry out this analysis on a "dummy" text fragment, which graphic designers refer to as the [_lorem ipsum_](https://en.wikipedia.org/wiki/Lorem_ipsum):

```python
latin_text = """
Sed ut perspiciatis, unde omnis iste natus error sit
voluptatem accusantium doloremque laudantium, totam
rem aperiam eaque ipsa, quae ab illo inventore
veritatis et quasi architecto beatae vitae dicta
sunt, explicabo. Nemo enim ipsam voluptatem, quia
voluptas sit, aspernatur aut odit aut fugit, sed
quia consequuntur magni dolores eos, qui ratione
voluptatem sequi nesciunt, neque porro quisquam est,
qui dolorem ipsum, quia dolor sit amet consectetur
adipisci[ng] velit, sed quia non numquam [do] eius
modi tempora inci[di]dunt, ut labore et dolore
magnam aliquam quaerat voluptatem. Ut enim ad minima
veniam, quis nostrum exercitationem ullam corporis
suscipit laboriosam, nisi ut aliquid ex ea commodi
consequatur? Quis autem vel eum iure reprehenderit,
qui in ea voluptate velit esse, quam nihil molestiae
consequatur, vel illum, qui dolorem eum fugiat, quo
voluptas nulla pariatur?

At vero eos et accusamus et iusto odio dignissimos
ducimus, qui blanditiis praesentium voluptatum
deleniti atque corrupti, quos dolores et quas
molestias excepturi sint, obcaecati cupiditate non
provident, similique sunt in culpa, qui officia
deserunt mollitia animi, id est laborum et dolorum
fuga. Et harum quidem rerum facilis est et expedita
distinctio. Nam libero tempore, cum soluta nobis est
eligendi optio, cumque nihil impedit, quo minus id,
quod maxime placeat, facere possimus, omnis voluptas
assumenda est, omnis dolor repellendus. Temporibus
autem quibusdam et aut officiis debitis aut rerum
necessitatibus saepe eveniet, ut et voluptates
repudiandae sint et molestiae non recusandae. Itaque
earum rerum hic tenetur a sapiente delectus, ut aut
reiciendis voluptatibus maiores alias consequatur
aut perferendis doloribus asperiores repellat.
"""

print("First 100 characters:\n  {} ...".format(latin_text[:100]))
```

Output:
```
First 100 characters:
  
Sed ut perspiciatis, unde omnis iste natus error sit
voluptatem accusantium doloremque laudantium,  ...
```

**Exercise 0** (ungraded). Look up and read the translation of _lorem ipsum_!

**Data cleaning.** Like most data in the real world, this dataset is noisy. It has both uppercase and lowercase letters, words have repeated letters, and there are all sorts of non-alphabetic characters. For our analysis, we should keep all the letters and spaces (so we can identify distinct words), but we should ignore case and ignore repetition within a word.

For example, the eighth word of this text is "error." As an _itemset_, it consists of the three unique letters, $\{e, o, r\}$. That is, treat the word as a set, meaning you only keep the unique letters.

This itemset has six possible _itempairs_: $\{e, o\}$, $\{e, r\}$, and $\{o, r\}$.

Start by writing some code to help "clean up" the input.

**Exercise 1** (`normalize_string_test`: 2 points). Complete the following function, `normalize_string(s)`, which should take a string (`str` object) `s` as input and returns a new string with all characters converted to lowercase and all non-alphabetic, non-space characters removed.

> Scanning the sample text, `latin_text`, you may see things that look like special cases. For instance, `inci[di]dunt` and `[do]`. For these, simply remove the non-alphabetic characters and only separate the words if there is explicit whitespace.
>
> For instance, `inci[di]dunt` would become `incididunt` (as a single word) and `[do]` would become `do` as a standalone word because the original string has whitespace on either side. A period or comma without whitespace would, similarly, just be treated as a non-alphabetic character inside a word _unless_ there is explicit whitespace. So `e pluribus.unum basium` would become `e pluribusunum basium` even though your common-sense understanding might separate `pluribus` and `unum`.

<!-- grade_id: normalize_string, points: 0, grade: false, solution: true -->
```python
import re
def normalize_string(s):
 '''   assert type (s) is str
    s=s.lower()
    print(s)
    regex = re.compile('[^a-zA-Z\s]')
    t=regex.sub('', s)

    return(t)'''
#
# YOUR CODE HERE
#
'''    initial=s.lower()
    final=""
    for i in initial:
        if i.isalpha()==True or i.isspace()==True:
            final = final+i
    return final '''
    
import re
def normalize_string(s):
    assert type (s) is str
    t=''
    s=s.lower()
    for i in s:
        if i.isalpha() or i.isspace():
            t=t+i
    return t

# Demo:
print(latin_text[:100], "...\n=>", normalize_string(latin_text[:100]), "...")
```

Output:
```
Sed ut perspiciatis, unde omnis iste natus error sit
voluptatem accusantium doloremque laudantium,  ...
=> 
sed ut perspiciatis unde omnis iste natus error sit
voluptatem accusantium doloremque laudantium  ...
```

<!-- grade_id: normalize_string_test, points: 2, grade: true, solution: false -->
```python
# `normalize_string_test`: Test cell
norm_latin_text = normalize_string(latin_text)

assert type(norm_latin_text) is str
assert len(norm_latin_text) == 1694
assert all([c.isalpha() or c.isspace() for c in norm_latin_text])
assert norm_latin_text == norm_latin_text.lower()

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 2** (`get_normalized_words_test`: 1 point). Implement the following function, `get_normalized_words(s)`: given a string (`str` object) `s`, returns a list of its words, normalized per the definition of `normalize_string()`.

<!-- grade_id: get_normalized_words, points: 0, grade: false, solution: true -->
```python
def get_normalized_words (s):
    assert type (s) is str
#
# YOUR CODE HERE
#
    '''initial=s.lower()
    final=""
    for i in initial:
        if i.isalpha()==True or i.isspace()==True:
            final = final+i
    return final.split() '''
    
    t=''
    s=s.lower()
    s=s.split()
    t=[i for i in s if i.isalpha() or i.isspace()]
    return t
# Demo:
print ("First five words:\n{}".format (get_normalized_words (latin_text)[:5]))

```

Output:
```
First five words:
['sed', 'ut', 'unde', 'omnis', 'iste']
```

<!-- grade_id: get_normalized_words_test, points: 1, grade: true, solution: false -->
```python
# `get_normalized_words_test`: Test cell
norm_latin_words = get_normalized_words(norm_latin_text)

assert len(norm_latin_words) == 250
for i, w in [(20, 'illo'), (73, 'eius'), (144, 'deleniti'), (248, 'asperiores')]:
    assert norm_latin_words[i] == w

print ("\n(Passed.)")
```

Output:
```
(Passed.)
```

**Exercise 3** (`make_itemsets_test`: 2 points). Implement a function, `make_itemsets`, that given a list of strings converts the characters of each string into an itemset, returning the list of itemsets.

<!-- grade_id: make_itemsets, points: 0, grade: false, solution: true -->
```python
def make_itemsets(words):
    l=[]
    for i in words:
        c=[]
        for j in i:
            c.append(j)
        itemset=set(c)
        l.append(itemset)
    
    return(l)
#
# YOUR CODE HERE
#
'''    itemset=[]
    for i in words:
        names=[]
        for j in i: 
            if j not in names:
                names.append(j)
        names=set(names)
        itemset.append(names)
            
    return itemset'''
    
words=['sed', 'ut', 'unde', 'omnis', 'iste']
print(make_itemsets(words))
```

Output:
```
[{'d', 'e', 's'}, {'u', 't'}, {'d', 'u', 'n', 'e'}, {'m', 'o', 's', 'n', 'i'}, {'e', 'i', 's', 't'}]
```

<!-- grade_id: make_itemsets_test, points: 2, grade: true, solution: false -->
```python
# `make_itemsets_test`: Test cell
norm_latin_itemsets = make_itemsets(norm_latin_words)

# Lists should have the same size
assert len(norm_latin_itemsets) == len(norm_latin_words)

# Test a random sample
from random import sample
for i in sample(range(len(norm_latin_words)), 5):
    print('[{}]'.format(i), norm_latin_words[i], "-->", norm_latin_itemsets[i])
    assert set(norm_latin_words[i]) == norm_latin_itemsets[i]
print("\n(Passed!)")
```

Output:
```
[20] illo --> {'o', 'i', 'l'}
[227] molestiae --> {'t', 'm', 'o', 's', 'i', 'l', 'a', 'e'}
[189] eligendi --> {'n', 'i', 'l', 'd', 'e', 'g'}
[22] veritatis --> {'r', 't', 'v', 's', 'i', 'a', 'e'}
[231] earum --> {'m', 'a', 'e', 'r', 'u'}

(Passed!)
```

## Implementing the basic algorithm

Recall the pseudocode for the algorithm that Rachel and Rich derived together:

![FindAssocRules (pseudocode)](https://ndownloader.figshare.com/files/7635700?private_link=3c473609741895a5cc2c)

In the following series of exercises, let's implement this method. We'll build it "bottom-up," first defining small pieces and working our way toward the complete algorithm. This method allows us to test each piece before combining them.

Observe that the bulk of the work in this procedure is just updating these tables, $T$ and $C$. So your biggest implementation decision is how to store those. A good choice is to use a dictionary

## Aside: Default dictionaries

Recall that the overall algorithm requires maintaining a table of item-pair (tuples) counts. It would be convenient to use a dictionary to store this table, where keys refer to item-pairs and the values are the counts.

However, with Python's built-in dictionaries, you always to have to check whether a key exists before updating it. For example, consider this code fragment:

```python
D = {'existing-key': 5} # Dictionary with one key-value pair

D['existing-key'] += 1 # == 6
D['new-key'] += 1  # Error: 'new-key' does not exist!
```

The second attempt causes an error because `'new-key'` is not yet a member of the dictionary. So, a more correct approach would be to do the following:

```python
D = {'existing-key': 5} # Dictionary with one key-value pair

if 'existing-key' not in D:
    D['existing-key'] = 0
D['existing-key'] += 1
   
if 'new-key' not in D:
    D['new-key'] = 0
D['new-key'] += 1
```

This pattern is so common that there is a special form of dictionary, called a _default dictionary_, which is available from the `collections` module: [`collections.defaultdict`](https://docs.python.org/3/library/collections.html?highlight=defaultdict#collections.defaultdict).

When you create a default dictionary, you need to provide a "factory" function that the dictionary can use to create an initial value when the key does *not* exist. For instance, in the preceding example, when the key was not present the code creates a new key with the initial value of an integer zero (0). Indeed, this default value is the one you get when you call `int()` with no arguments:

```python
print (int ())
```

Output:
```
0
```

```python
from collections import defaultdict

D2 = defaultdict (int) # Empty dictionary

D2['existing-key'] = 5 # Create one key-value pair

D2['existing-key'] += 1 # Update
D2['new-key'] += 1

print (D2)
```

Output:
```
defaultdict(<class 'int'>, {'new-key': 1, 'existing-key': 6})
```

**Exercise 4** (`update_pair_counts_test`: 2 points). Start by implementing a function that enumerates all item-pairs within an itemset and updates a table in-place that tracks the counts of those item-pairs.

You may assume all items in the given itemset (`itemset` argument) are distinct, i.e., that you may treat it as you would any set-like collection. You may also assume `pair_counts` is a default dictionary.

<!-- grade_id: update_pair_counts, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict
from itertools import combinations # Hint!

def update_pair_counts (pair_counts, itemset):
    """
    Updates a dictionary of pair counts for
    all pairs of items in a given itemset.
    """
    assert type (pair_counts) is defaultdict

#
# YOUR CODE HERE
#

    '''data=list(combinations(itemset, 2))
    data2=[(t[1], t[0]) for t in data]
    for i in data:
        pair_counts[i]+=1
    for j in data2:
        pair_counts[j]+=1
    return pair_counts '''
    
    for i,j in combinations(itemset,2):
        pair_counts[i,j]+=1
        pair_counts[j,i]+=1
    return pair_counts
    
```

<!-- grade_id: update_pair_counts_test, points: 2, grade: true, solution: false -->
```python
# `update_pair_counts_test`: Test cell
itemset_1 = set("error")
itemset_2 = set("dolor")
pair_counts = defaultdict(int)

update_pair_counts(pair_counts, itemset_1)
assert len(pair_counts) == 6
update_pair_counts(pair_counts, itemset_2)
assert len(pair_counts) == 16

print('"{}" + "{}"\n==> {}'.format (itemset_1, itemset_2, pair_counts))
for a, b in pair_counts:
    assert (b, a) in pair_counts
    assert pair_counts[(a, b)] == pair_counts[(b, a)]
    
print ("\n(Passed!)")
```

Output:
```
"{'o', 'e', 'r'}" + "{'d', 'o', 'r', 'l'}"
==> defaultdict(<class 'int'>, {('o', 'e'): 1, ('d', 'o'): 1, ('o', 'l'): 1, ('o', 'r'): 2, ('e', 'r'): 1, ('e', 'o'): 1, ('r', 'l'): 1, ('l', 'o'): 1, ('d', 'r'): 1, ('r', 'o'): 2, ('r', 'd'): 1, ('r', 'e'): 1, ('l', 'r'): 1, ('d', 'l'): 1, ('o', 'd'): 1, ('l', 'd'): 1})

(Passed!)
```

**Exercise 5** (`update_item_counts_test`: 2 points). Implement a procedure that, given an itemset, updates a table to track counts of each item.

As with the previous exercise, you may assume all items in the given itemset (`itemset`) are distinct, i.e., that you may treat it as you would any set-like collection. You may also assume the table (`item_counts`) is a default dictionary.

<!-- grade_id: update_item_counts, points: 0, grade: false, solution: true -->
```python
def update_item_counts(item_counts, itemset):
#
# YOUR CODE HERE
#
    '''for i in itemset:
        item_counts[i]+=1
    return item_counts'''
    for i in itemset:
        item_counts[i]+=1
    return item_counts
```

<!-- grade_id: update_item_counts_test, points: 2, grade: true, solution: false -->
```python
# `update_item_counts_test`: Test cell
itemset_1 = set("error")
itemset_2 = set("dolor")

item_counts = defaultdict(int)
update_item_counts(item_counts, itemset_1)
assert len(item_counts) == 3
update_item_counts(item_counts, itemset_2)
assert len(item_counts) == 5

assert item_counts['d'] == 1
assert item_counts['e'] == 1
assert item_counts['l'] == 1
assert item_counts['o'] == 2
assert item_counts['r'] == 2

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 6** (`filter_rules_by_conf_test`: 2 points). Given tables of item-pair counts and individual item counts, as well as a confidence threshold, return the rules that meet the threshold. The returned rules should be in the form of a dictionary whose key is the tuple, $(a, b)$ corresponding to the rule $a \Rightarrow b$, and whose value is the confidence of the rule, $\mathrm{conf}(a \Rightarrow b)$.

You may assume that if $(a, b)$ is in the table of item-pair counts, then both $a$ and $b$ are in the table of individual item counts.

<!-- grade_id: filter_rules_by_conf, points: 0, grade: false, solution: true -->
```python
def filter_rules_by_conf (pair_counts, item_counts, threshold):
    rules = {} # (item_a, item_b) -> conf (item_a => item_b)
#
# YOUR CODE HERE
#

    
    '''for (i,j) in pair_counts.keys():
        if item_counts[i]>0:
            chance=pair_counts[(i,j)]/item_counts[i]
            if chance>=threshold:
                rules[(i,j)]=chance
        
    return rules'''
    
    for (i,j) in pair_counts:
        assert i in item_counts
        conf_ab=pair_counts[(i,j)]/item_counts[i]
        if conf_ab>=threshold:
            rules[(i,j)]=conf_ab
        
    return rules




```

<!-- grade_id: filter_rules_by_conf_test, points: 2, grade: true, solution: false -->
```python
# `filter_rules_by_conf_test`: Test cell
pair_counts = {('man', 'woman'): 5,
               ('bird', 'bee'): 3,
               ('red fish', 'blue fish'): 7}
item_counts = {'man': 7,
               'bird': 9,
               'red fish': 11}
rules = filter_rules_by_conf (pair_counts, item_counts, 0.5)
print("Found these rules:", rules)
assert ('man', 'woman') in rules
assert ('bird', 'bee') not in rules
assert ('red fish', 'blue fish') in rules
print("\n(Passed!)")
```

Output:
```
Found these rules: {('man', 'woman'): 0.7142857142857143, ('red fish', 'blue fish'): 0.6363636363636364}

(Passed!)
```

**Aside: pretty printing the rules.** The output of rules above is a little messy; here's a little helper function that structures that output a little, which will be useful for both debugging and reporting purposes.

```python
def gen_rule_str(a, b, val=None, val_fmt='{:.3f}', sep=" = "):
    text = "{} => {}".format(a, b)
    if val:
        text = "conf(" + text + ")"
        text += sep + val_fmt.format(val)
    return text

def print_rules(rules):
    if type(rules) is dict or type(rules) is defaultdict:
        from operator import itemgetter
        ordered_rules = sorted(rules.items(), key=itemgetter(1), reverse=True)
    else: # Assume rules is iterable
        ordered_rules = [((a, b), None) for a, b in rules]
    for (a, b), conf_ab in ordered_rules:
        print(gen_rule_str(a, b, conf_ab))

# Demo:
print_rules(rules)
```

Output:
```
conf(man => woman) = 0.714
conf(red fish => blue fish) = 0.636
```

**Exercise 7** (`find_assoc_rules_test`: 3 points). Using the building blocks you implemented above, complete a function `find_assoc_rules` so that it implements the basic association rule mining algorithm and returns a dictionary of rules.

In particular, your implementation may assume the following:

1. As indicated in its signature, below, the function takes two inputs: `receipts` and `threshold`.
1. The input, `receipts`, is a collection of itemsets: for every receipt `r` in `receipts`, `r` may be treated as a collection of unique items.
2. The input `threshold` is the minimum desired confidence value. That is, the function should only return rules whose confidence is at least `threshold`.

The returned dictionary, `rules`, should be keyed by tuples $(a, b)$ corresponding to the rule $a \Rightarrow b$; each value should the the confidence $\mathrm{conf}(a \Rightarrow b)$ of the rule.

<!-- grade_id: find_assoc_rules, points: 0, grade: false, solution: true -->
```python
def find_assoc_rules(receipts, threshold):
#
# YOUR CODE HERE
#
    

    pair_counts=defaultdict(int)
    item_counts=defaultdict(int)
    for i in receipts:
        pair= update_pair_counts (pair_counts,i)
        print(pair)
        item=update_item_counts(item_counts, i)
        rule=filter_rules_by_conf (pair, item, threshold)
        
    return rule

```

<!-- grade_id: find_assoc_rules_test, points: 3, grade: true, solution: false -->
```python
# `find_assoc_rules_test`: Test cell
receipts = [set('abbc'), set('ac'), set('a')]
rules = find_assoc_rules(receipts, 0.6)

print("Original receipts as itemsets:", receipts)
print("Resulting rules:")
print_rules(rules)

assert ('a', 'b') not in rules
assert ('b', 'a') in rules
assert ('a', 'c') in rules
assert ('c', 'a') in rules
assert ('b', 'c') in rules
assert ('c', 'b') not in rules

print("\n(Passed!)")
```

Output:
```
defaultdict(<class 'int'>, {('c', 'a'): 1, ('b', 'a'): 1, ('a', 'c'): 1, ('a', 'b'): 1, ('b', 'c'): 1, ('c', 'b'): 1})
defaultdict(<class 'int'>, {('c', 'a'): 2, ('b', 'a'): 1, ('a', 'c'): 2, ('a', 'b'): 1, ('b', 'c'): 1, ('c', 'b'): 1})
defaultdict(<class 'int'>, {('c', 'a'): 2, ('b', 'a'): 1, ('a', 'c'): 2, ('a', 'b'): 1, ('b', 'c'): 1, ('c', 'b'): 1})
Original receipts as itemsets: [{'c', 'a', 'b'}, {'a', 'c'}, {'a'}]
Resulting rules:
conf(b => a) = 1.000
conf(b => c) = 1.000
conf(c => a) = 1.000
conf(a => c) = 0.667

(Passed!)
```

**Exercise 8** (`latin_rules_test`: 2 points). For the Latin string, `latin_text`, use your `find_assoc_rules()` function to compute the rules whose confidence is at least 0.75. Store your result in a variable named `latin_rules`.

<!-- grade_id: latin_rules, points: 0, grade: false, solution: true -->
```python
# Generate `latin_rules`:
#
# YOUR CODE HERE
#
a=get_normalized_words(latin_text)
b=make_itemsets(a)
latin_rules=find_assoc_rules(b,.75)
print_rules(latin_rules)
```

Output:
```
defaultdict(<class 'int'>, {('e', 'd'): 1, ('d', 's'): 1, ('s', 'e'): 1, ('d', 'e'): 1, ('e', 's'): 1, ('s', 'd'): 1})
defaultdict(<class 'int'>, {('e', 'd'): 1, ('t', 'u'): 1, ('d', 's'): 1, ('s', 'e'): 1, ('u', 't'): 1, ('d', 'e'): 1, ('e', 's'): 1, ('s', 'd'): 1})
defaultdict(<class 'int'>, {('e', 'd'): 2, ('d', 'u'): 1, ('t', 'u'): 1, ('u', 'n'): 1, ('s', 'e'): 1, ('n', 'u'): 1, ('u', 't'): 1, ('u', 'e'): 1, ('d', 'e'): 2, ('n', 'e'): 1, ('e', 'n'): 1, ('d', 's'): 1, ('n', 'd'): 1, ('d', 'n'): 1, ('s', 'd'): 1, ('u', 'd'): 1, ('e', 's'): 1, ('e', 'u'): 1})
defaultdict(<class 'int'>, {('i', 'm'): 1, ('e', 'd'): 2, ('n', 'o'): 1, ('s', 'n'): 1, ('n', 'i'): 1, ('n', 'm'): 1, ('u', 'd'): 1, ('d', 's'): 1, ('d', 'e'): 2, ('u', 'e'): 1, ('m', 'i'): 1, ('i', 'o'): 1, ('n', 'u'): 1, ('o', 'i'): 1, ('u', 'n'): 1, ('o', 's'): 1, ('e', 'u'): 1, ('i', 'n'): 1, ('d', 'u'): 1, ('m', 's'): 1, ('t', 'u'): 1, ('o', 'm'): 1, ('u', 't'): 1, ('s', 'i'): 1, ('s', 'd'): 1, ('m', 'o'): 1, ('n', 'e'): 1, 
```

<!-- grade_id: latin_rules_test, points: 2, grade: true, solution: false -->
```python
# `latin_rules_test`: Test cell
assert len(latin_rules) == 10
assert all([0.75 <= v <= 1.0 for v in latin_rules.values()])
for ab in ['xe', 'qu', 'hi', 'xi', 'vt', 're', 've', 'fi', 'gi', 'bi']:
    assert (ab[0], ab[1]) in latin_rules
print("\n(Passed!)")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-142-e1c34b7d584a> in <module>()
      3 assert all([0.75 <= v <= 1.0 for v in latin_rules.values()])
      4 for ab in ['xe', 'qu', 'hi', 'xi', 'vt', 're', 've', 'fi', 'gi', 'bi']:
----> 5     assert (ab[0], ab[1]) in latin_rules
      6 print("\n(Passed!)")

AssertionError:
```

As a final exercise, let's look at rules common to Latin text _and_ English text. You'll need the English translation of the _lorem ipsum_ text, encoded as the variable `english_text` in the next code cell:

```python
english_text = """
But I must explain to you how all this mistaken idea
of denouncing of a pleasure and praising pain was
born and I will give you a complete account of the
system, and expound the actual teachings of the great
explorer of the truth, the master-builder of human
happiness. No one rejects, dislikes, or avoids
pleasure itself, because it is pleasure, but because
those who do not know how to pursue pleasure
rationally encounter consequences that are extremely
painful. Nor again is there anyone who loves or
pursues or desires to obtain pain of itself, because
it is pain, but occasionally circumstances occur in
which toil and pain can procure him some great
pleasure. To take a trivial example, which of us
ever undertakes laborious physical exercise, except
to obtain some advantage from it? But who has any
right to find fault with a man who chooses to enjoy
a pleasure that has no annoying consequences, or
one who avoids a pain that produces no resultant
pleasure?

On the other hand, we denounce with righteous
indignation and dislike men who are so beguiled and
demoralized by the charms of pleasure of the moment,
so blinded by desire, that they cannot foresee the
pain and trouble that are bound to ensue; and equal
blame belongs to those who fail in their duty
through weakness of will, which is the same as
saying through shrinking from toil and pain. These
cases are perfectly simple and easy to distinguish.
In a free hour, when our power of choice is
untrammeled and when nothing prevents our being
able to do what we like best, every pleasure is to
be welcomed and every pain avoided. But in certain
circumstances and owing to the claims of duty or
the obligations of business it will frequently
occur that pleasures have to be repudiated and
annoyances accepted. The wise man therefore always
holds in these matters to this principle of
selection: he rejects pleasures to secure other
greater pleasures, or else he endures pains to
avoid worse pains.
"""
```

**Exercise 9** (`intersect_keys_test`: 2 points). Write a function that, given two dictionaries, finds the intersection of their keys.

<!-- grade_id: intersect_keys, points: 0, grade: false, solution: true -->
```python
def intersect_keys(d1, d2):
    assert type(d1) is dict or type(d1) is defaultdict
    assert type(d2) is dict or type(d2) is defaultdict
#
# YOUR CODE HERE
#
    '''a=set(d1.keys())
    b=set(d2.keys())
    return(a-(a-b))'''
    a=list(d1.keys())
    b=list(d2.keys())
    return [element for element in a if element in b]

d1={'s': 1, 'e': 1, 'd': 1}
d2={('c', 'b'): 1, ('a', 'b'): 1, ('c', 'a'): 2, ('b', 'a'): 1, ('a', 'c'): 2, ('b', 'c'): 1}
print(intersect_keys(d1,d2))
```

Output:
```
[]
```

<!-- grade_id: intersect_keys_test, points: 2, grade: true, solution: false -->
```python
# `intersect_keys_test`: Test cell
from random import sample

key_space = {'ape', 'baboon', 'bonobo', 'chimp', 'gorilla', 'monkey', 'orangutan'}
val_space = range(100)

for trial in range(10): # Try 10 random tests
    d1 = {k: v for k, v in zip(sample(key_space, 4), sample(val_space, 4))}
    d2 = {k: v for k, v in zip(sample(key_space, 3), sample(val_space, 3))}
    k_common = intersect_keys(d1, d2)
    for k in key_space:
        is_common = (k in k_common) and (k in d1) and (k in d2)
        is_not_common = (k not in k_common) and ((k not in d1) or (k not in d2))
        assert is_common or is_not_common
        
print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 10** (`common_high_conf_rules_test`: 1 points). Let's consider any rules with a confidence of at least 0.75 to be a "high-confidence rule."

Write some code that finds all high-confidence rules appearing in _both_ the Latin text _and_ the English text. Store your result in a list named `common_high_conf_rules` whose elements are $(a, b)$ pairs corresponding to the rules $a \Rightarrow b$.

<!-- grade_id: common_high_conf_rules, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
'''a=get_normalized_words(english_text)
b=make_itemsets(a)
english_rules=find_assoc_rules(b,.75)
common_high_conf_rules=intersect_keys(english_rules,latin_rules)'''

a=get_normalized_words(english_text)
b=make_itemsets(a)
englist_rules=find_assoc_rules(b,.75)
common_high_conf_rules=intersect_keys(english_rules,latin_rules)

print("High-confidence rules common to _lorem ipsum_ in Latin and English:")
print_rules(common_high_conf_rules)
```

Output:
```
defaultdict(<class 'int'>, {('b', 't'): 1, ('b', 'u'): 1, ('u', 'b'): 1, ('t', 'u'): 1, ('t', 'b'): 1, ('u', 't'): 1})
defaultdict(<class 'int'>, {('b', 't'): 1, ('b', 'u'): 1, ('u', 'b'): 1, ('t', 'u'): 1, ('t', 'b'): 1, ('u', 't'): 1})
defaultdict(<class 'int'>, {('u', 's'): 1, ('b', 'u'): 1, ('m', 's'): 1, ('b', 't'): 1, ('t', 'u'): 2, ('t', 'm'): 1, ('u', 't'): 2, ('t', 's'): 1, ('s', 'u'): 1, ('s', 't'): 1, ('m', 'u'): 1, ('u', 'b'): 1, ('u', 'm'): 1, ('t', 'b'): 1, ('s', 'm'): 1, ('m', 't'): 1})
defaultdict(<class 'int'>, {('p', 'e'): 1, ('x', 'l'): 1, ('l', 'e'): 1, ('b', 'u'): 1, ('x', 'a'): 1, ('b', 't'): 1, ('i', 'e'): 1, ('i', 'a'): 1, ('e', 'p'): 1, ('i', 'l'): 1, ('s', 'u'): 1, ('s', 't'): 1, ('e', 'a'): 1, ('x', 'i'): 1, ('u', 'm'): 1, ('p', 'a'): 1, ('m', 't'): 1, ('l', 'a'): 1, ('a', 'i'): 1, ('l', 'x'): 1, ('u', 's'): 1, ('l', 'i'): 1, ('p', 'x'): 1, ('t', 'u'): 2, ('n', 'p'): 1, ('n', 'e'): 1, ('t', 's'): 1, ('m', 'u'): 1, ('x', 'p'): 1, ('e', 'i'): 1, ('a', 'p'): 1, 
```

<!-- grade_id: common_high_conf_rules_test, points: 1, grade: true, solution: false -->
```python
# `common_high_conf_rules_test`: Test cell
assert len(common_high_conf_rules) == 2
assert ('x', 'e') in common_high_conf_rules
assert ('q', 'u') in common_high_conf_rules
print("\n(Passed!)")
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-108-5e53832ffee0> in <module>()
      1 # `common_high_conf_rules_test`: Test cell
----> 2 assert len(common_high_conf_rules) == 2
      3 assert ('x', 'e') in common_high_conf_rules
      4 assert ('q', 'u') in common_high_conf_rules
      5 print("\n(Passed!)")

NameError: name 'common_high_conf_rules' is not defined
```

## Putting it all together: Actual baskets!

Let's take a look at some real data that [someone](http://www.salemmarafi.com/code/market-basket-analysis-with-r/) was kind enough to prepare for a similar exercise designed for the R programming environment.

First, here's a code snippet to download the data, which is a text file:

```python
import requests
response = requests.get ('https://cse6040.gatech.edu/datasets/groceries.csv')
groceries_file = response.text  # or response.content for raw bytes

print (groceries_file[0:250] + "...\n... (etc.) ...") # Prints the first 250 characters only
```

Output:
```
citrus fruit,semi-finished bread,margarine,ready soups
tropical fruit,yogurt,coffee
whole milk
pip fruit,yogurt,cream cheese ,meat spreads
other vegetables,whole milk,condensed milk,long life bakery product
whole milk,butter,yogurt,rice,abrasive clea...
... (etc.) ...
```

Each line of this file is some customer's shopping basket. The items that the customer bought are stored as a comma-separated list of values.

**Exercise 11: Your task.** (`basket_rules_test`: 4 points). Your final task in this notebook is to mine this dataset for pairwise association rules. In particular, your code should produce (no pun intended!) a final dictionary, `basket_rules`, that meet these conditions (read carefully!):

1. The keys are pairs $(a, b)$, where $a$ and $b$ are item names (as strings).
2. The values are the corresponding confidence scores, $\mathrm{conf}(a \Rightarrow b)$.
3. Only include rules $a \Rightarrow b$ where item $a$ occurs at least `MIN_COUNT` times and $\mathrm{conf}(a \Rightarrow b)$ is at least `THRESHOLD`.

Pay particular attention to Condition 3: not only do you have to filter by a confidence threshold, but you must exclude rules $a \Rightarrow b$ where the item $a$ does not appear "often enough." There is a code cell below that defines values of `MIN_COUNT` and `THRESHOLD`, but your code should work even if we decide to change those values later on.

> _Aside_: Why would an analyst want to enforce Condition 3?

Your solution can use the `groceries_file` string variable defined above as its starting point. And since it's in the same notebook, you may, of course, reuse any of the code you've written above as needed. Lastly, if you feel you need additional code cells, you can create them _after_ the code cell marked for your solution but _before_ the code marked, `### TEST CODE ###`.

```python
# Confidence threshold
THRESHOLD = 0.5

# Only consider rules for items appearing at least `MIN_COUNT` times.
MIN_COUNT = 10
```

<!-- grade_id: basket_rules, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
normal=groceries_file.splitlines()

receipt=[]
for i in range(0,len(normal)):
    receipt.append(normal[i].split(','))

pair_counts=defaultdict(int)
item_counts=defaultdict(int)
for i in receipt:
    item=update_item_counts(item_counts, i)

item2 = { key:value for key, value in item.items() if value <10 }
keys=item2.keys()
for j in receipt:
    j=j-keys
    pair= update_pair_counts (pair_counts,j)
#item = dict((k, v) for k, v in item.items() if v >= MIN_COUNT)
basket_rules=filter_rules_by_conf (pair, item, THRESHOLD)


#item2 = { key:value for key, value in item.items() if value <10 }
#entries=set(item2.keys())

#print(basket_rules)
#print(basket_rules.keys())


#for key in entries:
        #if key in basket_rules.keys()[0]:
            #del basket_rules[key]
#for k,v in basket_rules.items():
    #print(k[0])


#print(basket_rules)



```

<!-- grade_id: basket_rules_test, points: 4, grade: true, solution: false -->
```python
### `basket_rules_test`: TEST CODE ###
print("Found {} rules whose confidence exceeds {}.".format(len(basket_rules), THRESHOLD))
print("Here they are:\n")
print_rules(basket_rules)

assert len(basket_rules) == 19
assert all([THRESHOLD <= v < 1.0 for v in basket_rules.values()])
ans_keys = [("pudding powder", "whole milk"), ("tidbits", "rolls/buns"), ("cocoa drinks", "whole milk"), ("cream", "sausage"), ("rubbing alcohol", "whole milk"), ("honey", "whole milk"), ("frozen fruits", "other vegetables"), ("cream", "other vegetables"), ("ready soups", "rolls/buns"), ("cooking chocolate", "whole milk"), ("cereals", "whole milk"), ("rice", "whole milk"), ("specialty cheese", "other vegetables"), ("baking powder", "whole milk"), ("rubbing alcohol", "butter"), ("rubbing alcohol", "citrus fruit"), ("jam", "whole milk"), ("frozen fruits", "whipped/sour cream"), ("rice", "other vegetables")]
for k in ans_keys:
    assert k in basket_rules

print("\n(Passed!)")
```

Output:
```
Found 19 rules whose confidence exceeds 0.5.
Here they are:

conf(honey => whole milk) = 0.733
conf(frozen fruits => other vegetables) = 0.667
conf(cereals => whole milk) = 0.643
conf(rice => whole milk) = 0.613
conf(rubbing alcohol => whole milk) = 0.600
conf(cocoa drinks => whole milk) = 0.591
conf(pudding powder => whole milk) = 0.565
conf(jam => whole milk) = 0.547
conf(cream => other vegetables) = 0.538
conf(cream => sausage) = 0.538
conf(baking powder => whole milk) = 0.523
conf(tidbits => rolls/buns) = 0.522
conf(cooking chocolate => whole milk) = 0.520
conf(rice => other vegetables) = 0.520
conf(frozen fruits => whipped/sour cream) = 0.500
conf(rubbing alcohol => butter) = 0.500
conf(rubbing alcohol => citrus fruit) = 0.500
conf(specialty cheese => other vegetables) = 0.500
conf(ready soups => rolls/buns) = 0.500

(Passed!)
```

```python

```