# Python review: More exercises

This notebook continues the review of Python basics based on [Chris Simpkins's](https://www.cc.gatech.edu/~simpkins/) [Python Bootcamp](https://www.cc.gatech.edu/~simpkins/teaching/python-bootcamp/syllabus.html).

This particular notebook adapts the exercises that appeared with the ["Functional Programming" slides](https://www.cc.gatech.edu/~simpkins/teaching/python-bootcamp/slides/functional-programming.html) of the Fall 2016 offering.

Consider the following dataset of exam grades, organized as a 2-D table and stored in Python as a "list of lists" under the variable name, `grades`.

```python
grades = [
    # First line is descriptive header. Subsequent lines hold data
    ['Student', 'Exam 1', 'Exam 2', 'Exam 3'],
    ['Thorny', '100', '90', '80'],
    ['Mac', '88', '99', '111'],
    ['Farva', '45', '56', '67'],
    ['Rabbit', '59', '61', '67'],
    ['Ursula', '73', '79', '83'],
    ['Foster', '89', '97', '101']
]
```

**Exercise 0** (`students_test`: 1 point). Write some code that computes a new list named `students[:]`, which holds the names of the students as they from "top to bottom" in the table.

<!-- grade_id: students, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
#students=[]
#assert type(students) is list
#for i in range(1,len(grades)):
 #   data=grades[i]
    #students.append(data[0])

#print(students)
#print (len(grades))

students=[i[0] for i in grades]
students=students[1:]
print(students)
```

Output:
```
['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']
```

<!-- grade_id: students_test, points: 1, grade: true, solution: false -->
```python
# `students_test`: Test cell
print(students)
assert type(students) is list
assert students == ['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']
print("\n(Passed!)")
```

Output:
```
['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']

(Passed!)
```

**Exercise 1** (`assignments_test`: 1 point). Write some code to compute a new list named `assignments[:]`, to hold the names of the class assignments. (These appear in the descriptive header element of `grades`.)

<!-- grade_id: assignments, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
assignments=[]
assert type(assignments) is list
assignments=grades[0][1:]
```

<!-- grade_id: assignments_test, points: 1, grade: true, solution: false -->
```python
# `assignments_test`: Test cell
print(assignments)
assert type(assignments) is list
assert assignments == ['Exam 1', 'Exam 2', 'Exam 3']
print("\n(Passed!)")
```

Output:
```
['Exam 1', 'Exam 2', 'Exam 3']

(Passed!)
```

**Exercise 2** (`grade_lists_test`: 1 point). Write some code to compute a new _dictionary_, named `grade_lists`, that maps names of students to _lists_ of their exam grades. The grades should be converted from strings to integers. For instance, `grade_lists['Thorny'] == [100, 90, 80]`.

<!-- grade_id: grade_lists, points: 0, grade: false, solution: true -->
```python
# Create a dict mapping names to lists of grades.
#
# YOUR CODE HERE
#
#grade_lists={}
#assert type(grade_lists) is dict
#name=[]
#score=[]
#for i in range(1,len(grades)):
 #   name.append(grades[i][0])
  #  score.append(grades[i][1:])
#score=[[int(y) for y in x] for x in score]

#for j in range(0,len(name)):
 #   student=name[j]
  #  grade_lists[student]=score[j]

#print(grade_lists)
names=[i[0] for i in grades ]
names=names[1:]
grades2=[i[1:] for i in grades]
grades2=grades2[1:]
for i,j in enumerate(grades2):
    grades2[i]=[int(k)for k in j]

print(names)
print(grades2)

grade_lists={}
for i in range(0,len(names)):
    grade_lists[names[i]]=grades2[i]
    
print(grade_lists)

```

Output:
```
['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']
[[100, 90, 80], [88, 99, 111], [45, 56, 67], [59, 61, 67], [73, 79, 83], [89, 97, 101]]
{'Thorny': [100, 90, 80], 'Foster': [89, 97, 101], 'Ursula': [73, 79, 83], 'Rabbit': [59, 61, 67], 'Farva': [45, 56, 67], 'Mac': [88, 99, 111]}
```

<!-- grade_id: grade_lists_test, points: 1, grade: true, solution: false -->
```python
# `grade_lists_test`: Test cell
print(grade_lists)
assert type(grade_lists) is dict, "Did not create a dictionary."
assert len(grade_lists) == len(grades)-1, "Dictionary has the wrong number of entries."
assert {'Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster'} == set(grade_lists.keys()), "Dictionary has the wrong keys."
assert grade_lists['Thorny'] == [100, 90, 80], 'Wrong grades for: Thorny'
assert grade_lists['Mac'] == [88, 99, 111], 'Wrong grades for: Mac'
assert grade_lists['Farva'] == [45, 56, 67], 'Wrong grades for: Farva'
assert grade_lists['Rabbit'] == [59, 61, 67], 'Wrong grades for: Rabbit'
assert grade_lists['Ursula'] == [73, 79, 83], 'Wrong grades for: Ursula'
assert grade_lists['Foster'] == [89, 97, 101], 'Wrong grades for: Foster'
print("\n(Passed!)")
```

Output:
```
{'Thorny': [100, 90, 80], 'Foster': [89, 97, 101], 'Ursula': [73, 79, 83], 'Rabbit': [59, 61, 67], 'Farva': [45, 56, 67], 'Mac': [88, 99, 111]}

(Passed!)
```

**Exercise 3** (`grade_dicts_test`: 2 points). Write some code to compute a new dictionary, `grade_dicts`, that maps names of students to _dictionaries_ containing their scores. Each entry of this scores dictionary should be keyed on assignment name and hold the corresponding grade as an integer. For instance, `grade_dicts['Thorny']['Exam 1'] == 100`.

<!-- grade_id: grade_dicts, points: 0, grade: false, solution: true -->
```python
# Create a dict mapping names to dictionaries of grades.
#
# YOUR CODE HERE
#
subject=grades[0][1:]


print (subject)
print(score)
print(name)
grade_dicts={}

assert type(grade_dicts) is dict

for i in range(0,len(name)):
    grade_dicts[name[i]]={}
    for j in range(0,len(subject)):
        grade_dicts[name[i]][subject[j]]=score[i][j]
        
print(grade_dicts)

        
```

Output:
```
['Exam 1', 'Exam 2', 'Exam 3']
[[100, 90, 80], [88, 99, 111], [45, 56, 67], [59, 61, 67], [73, 79, 83], [89, 97, 101]]
['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']
{'Foster': {'Exam 3': 101, 'Exam 2': 97, 'Exam 1': 89}, 'Ursula': {'Exam 3': 83, 'Exam 2': 79, 'Exam 1': 73}, 'Rabbit': {'Exam 3': 67, 'Exam 2': 61, 'Exam 1': 59}, 'Thorny': {'Exam 3': 80, 'Exam 2': 90, 'Exam 1': 100}, 'Farva': {'Exam 3': 67, 'Exam 2': 56, 'Exam 1': 45}, 'Mac': {'Exam 3': 111, 'Exam 2': 99, 'Exam 1': 88}}
```

<!-- grade_id: grade_dicts_test, points: 2, grade: true, solution: false -->
```python
# `grade_dicts_test`: Test cell
print(grade_dicts)
assert type(grade_dicts) is dict, "Did not create a dictionary."
assert len(grade_dicts) == len(grades)-1, "Dictionary has the wrong number of entries."
assert {'Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster'} == set(grade_dicts.keys()), "Dictionary has the wrong keys."
assert grade_dicts['Foster']['Exam 1'] == 89, 'Wrong score'
assert grade_dicts['Foster']['Exam 3'] == 101, 'Wrong score'
assert grade_dicts['Foster']['Exam 2'] == 97, 'Wrong score'
assert grade_dicts['Ursula']['Exam 1'] == 73, 'Wrong score'
assert grade_dicts['Ursula']['Exam 3'] == 83, 'Wrong score'
assert grade_dicts['Ursula']['Exam 2'] == 79, 'Wrong score'
assert grade_dicts['Rabbit']['Exam 1'] == 59, 'Wrong score'
assert grade_dicts['Rabbit']['Exam 3'] == 67, 'Wrong score'
assert grade_dicts['Rabbit']['Exam 2'] == 61, 'Wrong score'
assert grade_dicts['Mac']['Exam 1'] == 88, 'Wrong score'
assert grade_dicts['Mac']['Exam 3'] == 111, 'Wrong score'
assert grade_dicts['Mac']['Exam 2'] == 99, 'Wrong score'
assert grade_dicts['Farva']['Exam 1'] == 45, 'Wrong score'
assert grade_dicts['Farva']['Exam 3'] == 67, 'Wrong score'
assert grade_dicts['Farva']['Exam 2'] == 56, 'Wrong score'
assert grade_dicts['Thorny']['Exam 1'] == 100, 'Wrong score'
assert grade_dicts['Thorny']['Exam 3'] == 80, 'Wrong score'
assert grade_dicts['Thorny']['Exam 2'] == 90, 'Wrong score'
print("\n(Passed!)")
```

Output:
```
{'Foster': {'Exam 3': 101, 'Exam 2': 97, 'Exam 1': 89}, 'Ursula': {'Exam 3': 83, 'Exam 2': 79, 'Exam 1': 73}, 'Rabbit': {'Exam 3': 67, 'Exam 2': 61, 'Exam 1': 59}, 'Thorny': {'Exam 3': 80, 'Exam 2': 90, 'Exam 1': 100}, 'Farva': {'Exam 3': 67, 'Exam 2': 56, 'Exam 1': 45}, 'Mac': {'Exam 3': 111, 'Exam 2': 99, 'Exam 1': 88}}

(Passed!)
```

**Exercise 4** (`avg_grades_by_student_test`: 1 point). Write some code to compute a dictionary named `avg_grades_by_student` that maps each student to his or her average exam score. For instance, `avg_grades_by_student['Thorny'] == 90`.

> **Hint.** The [`statistics`](https://docs.python.org/3.5/library/statistics.html) module of Python has at least one helpful function.

<!-- grade_id: avg_grades_by_student, points: 0, grade: false, solution: true -->
```python
# Create a dict mapping names to grade averages.
#
# YOUR CODE HERE
#
from statistics import mean
avg_grades_by_student={}
assert type(avg_grades_by_student) is dict

for j in range(0,len(name)):
    student=name[j]
    avg_grades_by_student[student]=mean(score[j])
    
```

<!-- grade_id: avg_grades_by_student_test, points: 1, grade: true, solution: false -->
```python
# `avg_grades_by_student_test`: Test cell
print(avg_grades_by_student)
assert type(avg_grades_by_student) is dict, "Did not create a dictionary."
assert len(avg_grades_by_student) == len(students), "Output has the wrong number of students."
assert abs(avg_grades_by_student['Mac'] - 99.33333333333333) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Foster'] - 95.66666666666667) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Farva'] - 56) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Rabbit'] - 62.333333333333336) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Thorny'] - 90) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Ursula'] - 78.33333333333333) <= 4e-15, 'Mean is incorrect'
print("\n(Passed!)")
```

Output:
```
{'Foster': 95.66666666666667, 'Ursula': 78.33333333333333, 'Rabbit': 62.333333333333336, 'Thorny': 90, 'Farva': 56, 'Mac': 99.33333333333333}

(Passed!)
```

**Exercise 5** (`grades_by_assignment_test`: 2 points). Write some code to compute a dictionary named `grades_by_assignment`, whose keys are assignment (exam) names and whose values are lists of scores over all students on that assignment. For instance, `grades_by_assignment['Exam 1'] == [100, 88, 45, 59, 73, 89]`.

<!-- grade_id: grades_by_assignment, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
#grades_by_assignment_test={}
#assert type(grades_by_assignment_test) is dict
grades_by_ass=[]
grades_by_assignment={}
assert type(grades_by_assignment) is dict
subject=grades[0][1:]
for j in range(0, len(subject)):
    grades_by_ass.append([i[j] for i in score]) 

for j in range(0,len(subject)):
    sub=subject[j]
    grades_by_assignment[sub]=grades_by_ass[j]
    
print(grades_by_assignment)

```

Output:
```
{'Exam 3': [80, 111, 67, 67, 83, 101], 'Exam 2': [90, 99, 56, 61, 79, 97], 'Exam 1': [100, 88, 45, 59, 73, 89]}
```

<!-- grade_id: grades_by_assignment_test, points: 2, grade: true, solution: false -->
```python
# `grades_by_assignment_test`: Test cell
print(grades_by_assignment)
assert type(grades_by_assignment) is dict, "Output is not a dictionary."
assert len(grades_by_assignment) == 3, "Wrong number of assignments."
assert grades_by_assignment['Exam 1'] == [100, 88, 45, 59, 73, 89], 'Wrong grades list'
assert grades_by_assignment['Exam 3'] == [80, 111, 67, 67, 83, 101], 'Wrong grades list'
assert grades_by_assignment['Exam 2'] == [90, 99, 56, 61, 79, 97], 'Wrong grades list'
print("\n(Passed!)")
```

Output:
```
{'Exam 3': [80, 111, 67, 67, 83, 101], 'Exam 2': [90, 99, 56, 61, 79, 97], 'Exam 1': [100, 88, 45, 59, 73, 89]}

(Passed!)
```

**Exercise 6** (`avg_grades_by_assignment_test`: 1 point). Write some code to compute a dictionary, `avg_grades_by_assignment`, which maps each exam to its average score.

<!-- grade_id: avg_grades_by_assignment, points: 0, grade: false, solution: true -->
```python
# Create a dict mapping items to average for that item across all students.
#
# YOUR CODE HERE
#
from statistics import mean
avg_grades_by_assignment={}
assert type(avg_grades_by_assignment) is dict

for j in range(0,len(subject)):
    avg_grades_by_assignment[subject[j]]=mean(grades_by_ass[j])
    

print (avg_grades_by_assignment) 
```

Output:
```
{'Exam 3': 84.83333333333333, 'Exam 2': 80.33333333333333, 'Exam 1': 75.66666666666667}
```

<!-- grade_id: avg_grades_by_assignment_test, points: 1, grade: true, solution: false -->
```python
# `avg_grades_by_assignment_test`: Test cell
print(avg_grades_by_assignment)
assert type(avg_grades_by_assignment) is dict
assert len(avg_grades_by_assignment) == 3
assert abs((100+88+45+59+73+89)/6 - avg_grades_by_assignment['Exam 1']) <= 7e-15
assert abs((80+111+67+67+83+101)/6 - avg_grades_by_assignment['Exam 3']) <= 7e-15
assert abs((90+99+56+61+79+97)/6 - avg_grades_by_assignment['Exam 2']) <= 7e-15
print("\n(Passed!)")
```

Output:
```
{'Exam 3': 84.83333333333333, 'Exam 2': 80.33333333333333, 'Exam 1': 75.66666666666667}

(Passed!)
```

**Exercise 7** (`rank_test`: 2 points). Write some code to create a new list, `rank`, which contains the names of students in order by _decreasing_ score. That is, `rank[0]` should contain the name of the top student (highest average exam score), and `rank[-1]` should have the name of the bottom student (lowest average exam score).

<!-- grade_id: rank, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#
#total_grades_by_student={}
#assert type(total_grades_by_student) is dict
#for j in range(0,len(name)):
 #   student1=name[j]
  #  total_grades_by_student[student1]=sum(score[j])

#a=total_grades_by_student.keys()
#b=total_grades_by_student.values()
#c=list(zip(a,b))
#d=sorted(c,key=lambda x: x[1], reverse=True)
#rank=[item[0] for item in d]
#assert type(rank) is list

avg_grades_by_student={}
for i in grades[1:]:
    temp=[int(j) for j in i[1:]]
    num=0
    for j in temp:
        num=num+j
    
    avg_grades_by_student[i[0]]=num/3

    
a=list(avg_grades_by_student.keys())
b=list(avg_grades_by_student.values())
c=list(zip(a,b))
c.sort(key=lambda x: x[1])
d=[i[0] for i in c]
rank=d[::-1]

```

<!-- grade_id: rank_test, points: 2, grade: true, solution: false -->
```python
# `rank_test`: Test cell
print(rank)
print("\n=== Ranking ===")
for i, s in enumerate(rank):
    print("{}. {}: {}".format(i+1, s, avg_grades_by_student[s]))
    
assert rank == ['Mac', 'Foster', 'Thorny', 'Ursula', 'Rabbit', 'Farva']
for i in range(len(rank)-1):
    assert avg_grades_by_student[rank[i]] >= avg_grades_by_student[rank[i+1]]
print("\n(Passed!)")
```

Output:
```
['Mac', 'Foster', 'Thorny', 'Ursula', 'Rabbit', 'Farva']

=== Ranking ===
1. Mac: 99.33333333333333
2. Foster: 95.66666666666667
3. Thorny: 90.0
4. Ursula: 78.33333333333333
5. Rabbit: 62.333333333333336
6. Farva: 56.0

(Passed!)
```

```python

```